# 2 Boolean Algebra

### Task 1
**Find the missing values**

**a) d = False**

```python
a = False
b = ...
c = 0 < a
d = c or not a and not b
```

**b) g = True**
**&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;j = False**


```python
e = True
f = ...
g = 1 > f and e
h = ...
j = g and e and not f ^ h
```

Tests for Task 1.1 (points: None)
```
>>> assert d == False or d == 0
```
failure_message: Test failed. Keep the precedence in mind (NOT before AND before OR) and another hint, c is evaluated to False because a is interpreted as 0.


Tests for Task 1.2 (points: None)
```
>>> assert (j == False or j == 0) and (g == True or g == 1)
```
failure_message: Test failed. Keep the precedence in mind.
