# 4 Control Structures

### Task 1: Loops
**1.1 Write a while loop that iterates through the numbers from 0 to 100 and adds them to a list called enumerate_while.**


```python
enumerate_while = []
...
```

**1.2 Provide a second solution that uses a for loop and adds each iteration step to a list called enumerate_for.**


```python
enumerate_for = []
...
```

### Task 2
**Write code that iterates over `list1` and display numbers divisible by five. The code should stop the loop iteration if it finds a number greater than 150. Save the numbers divisible by 5 to `list_new`.**


```python
list1 = [5, 7, 12, 15, 32, 42, 55, 75, 91, 122, 132, 150, 157, 180, 181, 200]
...
```

Tests for Task 1.1 (points: None)
```
>>> assert enumerate_while == [0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30, 31, 32, 33, 34, 35, 36, 37, 38, 39, 40, 41, 42, 43, 44, 45, 46, 47, 48, 49, 50, 51, 52, 53, 54, 55, 56, 57, 58, 59, 60, 61, 62, 63, 64, 65, 66, 67, 68, 69, 70, 71, 72, 73, 74, 75, 76, 77, 78, 79, 80, 81, 82, 83, 84, 85, 86, 87, 88, 89, 90, 91, 92, 93, 94, 95, 96, 97, 98, 99, 100]
```
failure_message: Test failed. Remember not to exclude 100.


Tests for Task 1.2 (points: None)
```
>>> assert enumerate_for == [0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30, 31, 32, 33, 34, 35, 36, 37, 38, 39, 40, 41, 42, 43, 44, 45, 46, 47, 48, 49, 50, 51, 52, 53, 54, 55, 56, 57, 58, 59, 60, 61, 62, 63, 64, 65, 66, 67, 68, 69, 70, 71, 72, 73, 74, 75, 76, 77, 78, 79, 80, 81, 82, 83, 84, 85, 86, 87, 88, 89, 90, 91, 92, 93, 94, 95, 96, 97, 98, 99, 100]
```
failure_message: Test failed. Keep in mind, that the stop-Parameter of the range()-Function is not included.


Tests for Task 2.1 (points: None)
```
>>> assert list_new == [5, 15, 55, 75, 150]
```
failure_message: Test failed.
