# 5 Functions and Modules

### Task 1
**Write a function `square_fun` that returns the square of its argument.**

```python
def square_fun(x):
    ...
```

### Task 2
**Write a function `reverse_list` that takes a list of integers as input and returns the list in reverse.**

```python
def reverse_list(lst):
    ...
```

### Task 3
 **Write a function `is_prime` that takes an integer as input and returns `True` if it is a prime number, and `False` otherwise.**

```python
def is_prime(num):
    ...
```

### Task 4
**Write a function `is_even` that takes an integer as input and returns `True` if it is an even number, and `False` otherwise. \
Negative numbers are not considered even or odd. When entering a negative number the function should return `0`.**


```python
def is_even(num):
    ...
```

### Task 5
**Write a function `get_common_elements` that takes two lists as input and returns a new list containing the common elements between them.**

*Hint: You might want to use the set() function to transform the list to a set for the comparison. https://python-reference.readthedocs.io/en/latest/docs/sets/*

```python
def get_common_elements(lst1,lst2):
    ...
```

### Task 6
**Write a function `broken_by_7` that takes a list of numbers as input and finds the first occurence of a number that is divisible by 7. When such a number is found, the program should stop iterating and `return` the number. If such a number is not found the function should return `False`.**

```python
def broken_by_7(my_list):
    ...
```

### Task 7
**Rewrite the function `summarize` using a while loop. The function takes a positive integer as an input and then sums up all integers between 0 and the given number.**
*Hint: The autograder is currently unable to detect if you used the correct loop type.*

```python
def summarize(n):
    result = 0
    for i in range(0, n+1):
        result += i
    return result

# rewrite the above function into a while loop.

```

Tests for Task 1 Square Function (points: None)
```
>>> result = square_fun(2)
>>> assert result == 4
```
failure_message: Test failed at positive numbers
```
>>> result = square_fun(1)
>>> assert result == 1
```
failure_message: Test failed at 1
```
>>> result = square_fun(0)
>>> assert result == 0
```
failure_message: Test failed at 0
```
>>> result = square_fun(-3)
>>> assert result == 9
```
failure_message: Test failed at negative numbers


Tests for Task 2 Reverse List (points: None)
```
>>> my_list = [1,2,3,4]
>>> result = reverse_list(my_list)
>>> assert result == [4,3,2,1]
```
failure_message: List is not sorted correctly


Tests for Task 3 Prime Number Checker (points: None)
```
>>> result = is_prime(13)
>>> assert result == True
```
failure_message: Failed to classify a prime number
```
>>> result = is_prime(89)
>>> assert result == True
```
failure_message: Failed to classify a prime number
```
>>> result = is_prime(1)
>>> assert result == False
```
failure_message: Classified 1 as a prime number
```
>>> result = is_prime(-13)
>>> assert result == False
```
failure_message: Negative numbers are no prime numbers
```
>>> result = is_prime(4)
>>> assert result == False
```
failure_message: Classified a non-prime number as prime number
```
>>> result = is_prime(200)
>>> assert result == False
```
failure_message: Classified a non-prime number as prime number


Tests for Task 4 Odd or Even (points: None)
```
>>> result = is_even(10)
>>> assert result == True
```
failure_message: Failed to detect even number
```
>>> result = is_even(5)
>>> assert result == False
```
failure_message: Failed to detect uneven number
```
>>> result = is_even(-40)
>>> assert result == 0
```
failure_message: Function should return "Error" when entering negative number
success_message: Succesfully detected negative numbers as invalid input


Tests for Task 5 List Duplicate Checker (points: None)
```
>>> result = get_common_elements([1,2,3,4],[3,4,5,6])
>>> assert result == [3,4]
```
failure_message: Failed to get common items of the lists
```
>>> result = get_common_elements([],[])
>>> assert result == []
```
failure_message: Failed at empty lists
```
>>> result = get_common_elements([1,2,3],[7,8,9])
>>> assert result == []
```
failure_message: Failed at lists with no match
```
>>> result = get_common_elements([1,2,3],[7,8,9,3])
>>> assert result == [3]
```
failure_message: Failed at lists with different lengths
```
>>> result = get_common_elements([1,2,3,"a"],[7,8,9,"3","a"])
>>> assert result == ["a"]
```
failure_message: Failed at lists with different data types
```
>>> result = get_common_elements([1,2,3],[])
>>> assert result == []
```
failure_message: Failed at comparrison of empty list and filled list


Tests for Task 6 Break Statement (points: None)
```
>>> my_list = [1,2,3,4,5,6,7,8,9,14]
>>> assert broken_by_7(my_list) == 7
```
failure_message: Failed to detect the right number
success_message: Result correct
```
>>> my_list = []
>>> assert broken_by_7(my_list) == False
```
failure_message: Test failed for empty lists
```
>>> my_list = [14]
>>> assert broken_by_7(my_list) == 14
```
failure_message: Test failed for list with one entry
```
>>> my_list = [1,2,3,4,5,6,12,8,9,23]
>>> assert broken_by_7(my_list) == False
```
failure_message: Test failed for lists without valid numbers


Tests for Task 7 Transforming Loops (points: None)
```
>>> assert summarize(5) == 15
```
failure_message: Your solution produced a wrong sum
success_message: Your function output is correct. But did you use a while loop?
```
>>> assert summarize(10) == 55
```
failure_message: Your solution produced a wrong sum
success_message: Your function output is correct. But did you use a while loop?
```
>>> assert summarize(10) == 45
```
success_message: Check the condition of your while loop
