# 1 Language Elements

# Task 1: Operations
**Fill in the missing variables to complete the following calculations.**

**a) i = 4**

```python
a = ...
i = a + a
print(i)
```

**b) ii = 6.0**

```python
a = 5
b = ...
ii = a + b
print(ii)
```

**c) iii = 4.0**

```python
a = 2
b = 2.0
iii = ...
print(iii)
```

**d) iv = 2**

```python
c = ...
a = 1
iv = c - a
print(iv)
```

**e) v = 2.0**

```python
c = 3
a = ...
v = c - a
print(v)
```

**f) vi = 1.5**

```python
c = 6
a = 4
vi = ...
print(vi)
```

**g) vii = 1**

```python
a = 2
c = ...
vii = c // a
print(vii)

```

**h) viii = 4.0**

```python
b = ...
c = 4
viii = c / b
print(viii)
```

**i) ix = 2.0**

```python
a = 2
c = ...
ix = a // c
print(ix)
```

**j) x = 1**

```python
a = 2
b = ...
x = b % a
print(x)
```

**k) xi = 0.0**

```python
a = ...
c = 3
xi = c % a
print(xi)
```

**l) xii = 4**

```python
c = ...
xii = c * c
print(xii)
```

**m) xiii = 8.0**

```python
a = 8
c = ...
xiii = a * c
print(xiii)
```

**n) xiv = 8**

```python
a = 2
c = 3
xiv = ...
print(xiv)
```

**o) xv = 4.0**

```python
b = 4
a = ...
xv = b ** a
print(xv)
```

**p) xvi = 3**

```python
a = ...
b = 1.5
xvi = a + int(b)
print(xvi)
```

**q) xvii = 5**

```python
a = ...
b = 1.5
xvii = a + round(b)
```

**r) xviii = 6.0**


```python
a = ...
b = 1.5
xviii = a + float(int(b))

```

**s) xix = 9.0**

```python
a = ...
c = 3
xix = a + float(c)
```

# Task 2: Strings
**2.1 How are the following statements generated?**

**a) xx = 'Lebensmittelinformatik'**

```python
d = 'Lebensmittel'
e = ...
xx = d + e
print(xx)
```

**b) xxi = 'LebensmittelLebensmittel'**

```python
k = ...
xxi = 2 * k
print(xxi)
```

**c) xxii = 'forklift'**

```python
e = 'informatik'
f = e[2:5] + e[9]
g = ...
xxii = f + g
```

**d) xxiii = 'Vertical farming'**

```python
e = 'informatik'
h = e[2] + e[6] +e[4] + e[5]
j = ...
xxiii = 'Vertical ' + h + j
```

**2.2 What is the output of the following print() Functions?**

**d = 'Lebensmittel'**
**e = 'informatik'**

**a) print("d" + 'e')**

```python
output_1 = ...
```

**b) print(d + " " + e)**

```python
output_2 = ...
```

**c) print(d)**

```python
output_3 = ...
```

**2.3 Generate the word 'farm2fork' and save it using variable `k`.**


```python
h = 'fork'
j = 'farm'
k = ...
print(k)
```

**2.4 Apply the correct operations.**

_For example:_
_Task:&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;Count the a's appearing within the string l. Save the result using the variable `result`._

_Solution: &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;l = 'farm2fork'_
&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;_result = l.count('a')_

**a) Count the 0's appearing within the string l. Save the result using the variable `count_1`.**

```python
l = 'farm2fork'
count_1 = ...
```

**b) Count the o's appearing within the string l? Save the result using the variable `count_2`.**

```python
l = 'farm2fork'
count_2 = ...
```

**c) Check the length of the string l. Save the result using the variable `length_1`.**

```python
l = 'farm2fork'
length_1 = ...
```

### Task 3
**Use the string `truth` for the following tasks:**
1. Create a variable called `doubled_truth` which contains the string 2 times.
2. Create a variable called `reversed_truth` which contains the string in reversed order.
3. Create a variable called `half_truth` which contains exactly the first half of the string.
4. Create a variable called `connected_truth` which contains the contents of task 1,2  and 3.

```python
truth = "Foodinformatics is Funinformatics"

doubled_truth = ...
reversed_truth = ...
half_truth = ...
connected_truth = ...
```

Tests for Task 1 a Operations (points: None)
```
>>> assert i == 4 and isinstance(i, int)
```
failure_message: Test failed. Check, whether your variable is an integer.


Tests for Task 1 b Operations (points: None)
```
>>> assert ii == 6 and isinstance(ii, float)
```
failure_message: Test failed. Check, whether your variable is a float.


Tests for Task 1 c Operations (points: None)
```
>>> assert iii == 4 and isinstance(iii, float)
```
failure_message: Test failed. Try a summation.


Tests for Task 1 d Operations (points: None)
```
>>> assert iv == 2 and isinstance(iv, int)
```
failure_message: Test failed. Check, whether your variable is an integer.


Tests for Task 1 e Operations (points: None)
```
>>> assert v == 2 and isinstance(v, float)
```
failure_message: Test failed. Check, whether your variable is a float.


Tests for Task 1 f Operations (points: None)
```
>>> assert vi == 1.5
```
failure_message: Test failed. Try a division.


Tests for Task 1 g Operations (points: None)
```
>>> assert vii == 1 and isinstance(vii, int)
```
failure_message: Test failed. Check, whether your variable is an integer.


Tests for Task 1 h Operations (points: None)
```
>>> assert viii == 4 and isinstance(viii, float)
```
failure_message: Test failed. Check, whether your variable is a float.


Tests for Task 1 i Operations (points: None)
```
>>> assert ix == 2 and isinstance(ix, float)
```
failure_message: Test failed. Check, whether your variable is a float.


Tests for Task 1 j Operations (points: None)
```
>>> assert x == 1 and isinstance(x, int)
```
failure_message: Test failed. Check, whether your variable is an integer.


Tests for Task 1 k Operations (points: None)
```
>>> assert xi == 0 and isinstance(xi, float)
```
failure_message: Test failed. Check, whether your variable is a float.


Tests for Task 1 l Operations (points: None)
```
>>> assert xii == 4 and isinstance(xii, int)
```
failure_message: Test failed. Check, whether your variable is an integer.


Tests for Task 1 m Operations (points: None)
```
>>> assert xiii == 8 and isinstance(xiii, float)
```
failure_message: Test failed. Check, whether your variable is a float.


Tests for Task 1 n Operations (points: None)
```
>>> assert xiv == 8 and isinstance(xiv, int)
```
failure_message: Test failed. Try an exponentiation.


Tests for Task 1 o Operations (points: None)
```
>>> assert xv == 4 and isinstance(xv, float)
```
failure_message: Test failed. Check, whether your variable is a float.


Tests for Task 1 p Operations (points: None)
```
>>> assert xvi == 3 and isinstance(xvi, int)
```
failure_message: Test failed. Check, whether your variable a is an integer.


Tests for Task 1 q Operations (points: None)
```
>>> assert xvii == 5 and isinstance(xvii, int)
```
failure_message: Test failed. Check, whether your variable a is an integer.


Tests for Task 1 r Operations (points: None)
```
>>> assert xviii == 6 and isinstance(xviii, float)
```
failure_message: Test failed. Keep in mind, that float(int(b)) leads to 1.0


Tests for Task 1 s Operations (points: None)
```
>>> assert xix == 9 and isinstance(xix, float)
```
failure_message: Test failed.


Tests for Task 2.1 a  Operations (points: None)
```
>>> assert xx == 'Lebensmittelinformatik' and isinstance(xx, str)
```
failure_message: Test failed. Check, whether your variable e is a string.


Tests for Task 2.1 b  Operations (points: None)
```
>>> assert xxi == 'LebensmittelLebensmittel' and isinstance(xxi, str)
```
failure_message: Test failed. Check, whether your variable k is a string.


Tests for Task 2.1 c  Operations (points: None)
```
>>> assert xxii == 'forklift' and isinstance(xxii, str)
```
failure_message: Test failed. Check, whether your variable g is a string.


Tests for Task 2.1 d  Operations (points: None)
```
>>> assert xxiii == 'Vertical farming' and isinstance(xxiii, str)
```
failure_message: Test failed. Check, whether your variable j is a string.


Tests for Task 2.2 a  Strings (points: None)
```
>>> assert output_1 == 'de' and isinstance(output_1, str)
```
failure_message: Test failed. Check, whether your output_1 is a string.


Tests for Task 2.2 b   Strings (points: None)
```
>>> assert output_2 == 'Lebensmittel informatik' and isinstance(output_2, str)
```
failure_message: Test failed. Check, whether your output_2 is a string.


Tests for Task 2.2 c  Strings (points: None)
```
>>> assert output_3 == 'Lebensmittel' and isinstance(output_3, str)
```
failure_message: Test failed. Check, whether your output_3 is a string.


Tests for Task 2.3  Strings (points: None)
```
>>> assert k == 'farm2fork'
```
failure_message: Test failed. Keep in mind, that the '2' needs to be a string as well.


Tests for Task 2.4 a  Strings (points: None)
```
>>> assert count_1 == 0
```
failure_message: Test failed. Keep in mind, that the '0' needs to be a string.


Tests for Task 2.4 b  Strings (points: None)
```
>>> assert count_2 == 1
```
failure_message: Test failed. Keep in mind to only count the lower case letters.


Tests for Task 2.4 c Strings (points: None)
```
>>> assert length_1 == 9
```
failure_message: Test failed.


Tests for Task 3 (points: None)
```
>>> assert doubled_truth == 'Foodinformatics is FuninformaticsFoodinformatics is Funinformatics'
```
failure_message: Test failed for reversed_truth
```
>>> assert reversed_truth == 'scitamrofninuF si scitamrofnidooF'
```
failure_message: Test failed for reversed_truth
```
>>> assert half_truth == 'Foodinformatics ' or half_truth == 'Foodinformatics i'
```
failure_message: Test failed for half_truth - did you choose the correct point to split the string in half?
```
>>> assert connected_truth == 'Foodinformatics is FuninformaticsFoodinformatics is FuninformaticsscitamrofninuF si scitamrofnidooFFoodinformatics ' or connected_truth == 'Foodinformatics is FuninformaticsFoodinformatics is FuninformaticsscitamrofninuF si scitamrofnidooFFoodinformatics i'
```
failure_message: Test failed for connected_truth - did your solve all other tasks?
