# 3 Python Data Structures - List

### Task 1
**1.1 Creat a list my_fruits with the following items (use the given order): apple, cherry, peach, banana and strawberry.**

```python
...
```

**1.2 Append 'carrot' to the list my_vegetables.**

```python
my_vegetables = ['cauliflower', 'broccoli', 'cabbage', 'onion']
...
```

**1.3 Remove the rose from the list my_trees.**

```python
my_trees = ['birch', 'oak', 'maple', 'rose', 'spruce']
...
```

**1.4 Add basil between sage and mint to the list my_herbs.**

```python
my_herbs = ['thyme', 'dill', 'sage', 'mint', 'rosemary']
...
```

Tests for Task 1.1 (points: None)
```
>>> assert my_fruits == ['apple', 'cherry', 'peach', 'banana', 'strawberry']
```
failure_message: Test failed. Keep in mind, that the list items are applied as strings.


Tests for Task 1.2 (points: None)
```
>>> assert my_vegetables == ['cauliflower', 'broccoli', 'cabbage', 'onion', 'carrot']
```
failure_message: Test failed. Keep in mind, to apply 'carrot' as a string.


Tests for Task 1.3 (points: None)
```
>>> assert my_trees == ['birch', 'oak', 'maple', 'spruce']
```
failure_message: Test failed. Keep in mind, to apply 'rose' as a string.


Tests for Task 1.4 (points: None)
```
>>> assert my_herbs == ['thyme', 'dill', 'sage', 'basil', 'mint', 'rosemary']
```
failure_message: Test failed. Did you insert 'basil' at position 3?
